```{image} https://www.epfl.ch/about/overview/wp-content/uploads/2020/07/logo-epfl-1024x576.png 
:width: 140px
:align: left
```
## Image Processing Laboratory Notebooks
---

This Juypter notebook is part of a series of computer laboratories which are designed
to teach image-processing programming; they are running on the EPFL's Noto server. They are the practical complement of the theoretical lectures of the EPFL's Master course 
[**MICRO-511 Image Processing I**](https://moodle.epfl.ch/course/view.php?id=522) taught by Prof. M. Unser and Prof. D. Van de Ville.

The project is funded by the Center for Digital Education and the School of Engineering. It is owned by the [Biomedical Imaging Group](http://bigwww.epfl.ch/). 
The distribution or the reproduction of the notebook is strictly prohibited without the written consent of the authors.  &copy; EPFL 2023.

**Authors**: 
    [Pol del Aguila Pla](mailto:pol.delaguilapla@epfl.ch), 
    [Kay Lächler](mailto:kay.lachler@epfl.ch),
    [Alejandro Noguerón Arámburu](mailto:alejandro.nogueronaramburu@epfl.ch),
    [Yan Liu](mailto:yan.liu@epfl.ch), and
    [Daniel Sage](mailto:daniel.sage@epfl.ch).
    
---
# Lab 0: Introduction
**Released**: Thursday September 21, 2023

Double-click on this cell, fill your name and SCIPER number above, and run the cell below to verify your identity in Noto and set the seed for random results.

<!-- grade_id: cell-136865f50e6e2d19_2022, points: 0, grade: true, solution: false -->
```python
import getpass
# This line recovers your camipro number to mark the images with your ID
uid = int(getpass.getuser().split('-')[2]) if len(getpass.getuser().split('-')) > 2 else ord(getpass.getuser()[0])
print(f'SCIPER: {uid}')
```

# Introduction 

This introductory lab will guide you through the tools that you will need for the labs. 

## Table of contents
1. [Introduction to Jupyter notebooks](#intro-to-notebooks) 
2. [Elements of Python](#python)
    1. [Basic Python programming](#basic-python)
    2. [NumPy](#numpy)
    3. [Matplotlib](#matplotlib)
3. [Python Image Viewer](#python-image-viewer)
    1. [Creation of a viewer](#creation-of-a-viewer)
    2. [Using widgets](#using-widgets) 
    3. [User-defined widgets](#user-defined-widgets)
    4. [Programmatic customization](#programmatic-customization)
    5. [Try it yourself!](#3.5-try-it-yourself) 
4. [Image processing in Python](#image-processing-in-python)
    1. [Examples](#examples)
    2. [Try it yourself!](#4.2-try-it-yourself) 


(intro-to-notebooks)=
# 1. Introduction to Jupyter notebooks

[Back to table of contents](#Table-of-contents)

A Jupyter notebook consists of cells that contain either code snippets with their results or text. The cellular structure of Jupyter notebooks allows us to run small snippets of code one after the other, go back and forth between, or check intermediate results between them. 

<b>Note</b>: If you declare some variable in a cell, this variable will be available in all other cells that use the same programming language, irrespective of the location of the cell. Let's look at an example. Read the description of the cells below and run them as instructed.

In this first cell we declare a variable `a` and assign it the value $5$. Run it.

```python
# cell 1
# this is our variable
a = 5
```

In the second cell we print the value of `a`. If you run it just after the cell above, it should print "The value of a is 5". Run it.

```python
# cell 2
# here we print the value of a
print(f"The value of a is {a}")
```

Then, in the third cell, we modify `a` by adding $2$ to it. Run it.

```python
# cell 3
# add two to a
a = a + 2
```

Now run _cell 2_ again: the value of `a` has been modified, even though _cell 2_ is located above _cell 3_ in our notebook. Now run _cell 1_ again, where we set `a=5`, and `a` will be $5$ in all cells again. 
This shows the **global** impact every single cell has on the notebook. 

<div class="alert alert-info">

**Note:** If at any point the kernel should crash for whatever reason, you can restart it by clicking on the button `Kernel` in the toolbar on top and then click on `Restart Kernel...`. Once the kernel is restarted, you will need to run all necessary cells again to regain the working environment you had before. A good practice to test that there are no errors in the notebook is to click `Restart and Run All Cells...` once you have finished a notebook. This will run all the cells from top to bottom and only stop if an error occurs.
</div>

(python)=
# 2. Elements of Python

[Back to table of contents](#Table-of-contents)

_Python_ is a high-level language with many great libraries with built-in image processing algorithms. You will learn how to apply the different image processing methods to images using the fastest and most standard tools and visualize the results. 

    
**Note:** If you are already familiar with [Python](python.org) and the [NumPy](numpy.org) and [Matplotlib](matplotlib.org) libraries, you can jump to the next section.
</div>

(basic-python)=
## 2.1. Basic Python programming
[Back to table of contents](#Table-of-contents)

There are three basic types of variables  in Python:
```python
# Numbers
my_var = 2
my_var = 5.2
# Strings
my_var = 'Either use single quotes'
my_var = "or use double quotes"
# Boolean
my_var = True
my_var = False
```
Every variable can hold every type of data, so we don't need to specify the type at initialization.

We can have lists of several items:
```python
# List of numbers
my_list = [2.5, 7, 3.3]
# List of strings
my_list = ['E.T.', 'phone', 'home']
# List of booleans
my_list = [True, True, False]
```
Lists are 0-indexed and can be accessed with square brackets:
```python
# Set the 3rd element of my_list to 5
my_list[2] = 5
# Retrieve the 2nd element of my_list
new_var = my_list[1]
```
Similarly we can make tuples, which are just like lists, but their values cannot be modified after they have been initialized:
```python
# This is a tuple of three numbers
my_tuple = (3, 8.6, 4)
# Access to a tuple is similar to lists
new_var = my_tuple[1]
# But we cannot write to a tuple!
```
For numbers we have the following basic operators:
```python
# Define the two variables
a = 5
b = 2
# Addition
a + b
>> 7
# Subtraction
a - b
>> 3
# Multiplication
a * b
>> 10
# Division
a / b
>> 2.5
# Integer division
a // b
>> 2
# Modulo (remainder of integer division)
a % b
>> 1
# Power
a ** b
>> 25
```
We can have conditional statements
```python
if a < b:
    # printing is done with the print function
    print('a is smaller than b')
else:
    print('a is bigger than or equal to b')
```
<div class="alert alert-info">

**Note:** Indentation is important in Python. Writing a conditional statement, a loop, a function, etc., without an indent (`Tab`) raises an error.
</div>

with the following comparison operators:
```python
<  # Smaller than
<= # Smaller or equal
>  # Bigger than
>= # Bigger than or equal to
== # Equal to
!= # Not equal to
```
A `for` loop is declared as:
```python
# Note: range(n) creates numbers from 0 to n-1
for i in range(10):
    # Another for loop inside the for loop - we call this 'nested for loops'
    for j in range(10):
        # We can also have formatted strings indicated by a preceeding 'f', where the number i is inserted into the string
        print(f'Loop iteration number {j}')
```

Similarly a `while` loop:
```python
i = 0
while i < 10:
    print(f'Loop iteration number {i}')
    i = i + 1
```
And finally, a function can be defined using the `def` keyword:
```python
# A function that adds two numbers and returns the result
def addition(a, b):
    # Add the numbers
    result = a + b
    # Return the result
    return result

# Run the function
addition(5, 10)
>> 15
```

Experiment with Python in the cell below. Learn to do the simple things that come to mind and take the chance to explore all of the options above.

```python

```

(numpy)=
### 2.2. NumPy
[Back to table of contents](#Table-of-contents)

[NumPy](https://NumPy.org/) library provides multidimensional matrix/array processing tools. Most image processing libraries in Python use NumPy arrays to represent images. 

It is best practice to import the NumPy library as `np`, to avoid typing `numpy` every time. This also makes the code more readable. Run the cell below to import NumPy.

```python
# Import the NumPy library
import numpy as np
```

To create a new NumPy array filled with zeros we can use `np.zeros`. It requires the size of the array we want to get as a tuple `(height, width)`. Run the cell below to see an example.

Optionally, we can specify the internal numerical [type](https://NumPy.org/devdocs/user/basics.types.html) of the array. For example, one can use 8-bit unsigned integers (`np.uint8`). This type is convenient for 8-bit images, as the resulting array will occupy less memory than with other types, such as `np.float32` or `np.float64`. For most of the labs, you will not need to mind the numerical type of your variables: it adapts automatically to represent what you are storing in each instance.

```python
# A NumPy array of size 5x10 with numerical type "unsigned 8-bit integer"
my_arr = np.zeros((5,10), dtype=np.uint8)
# Print the result
print(my_arr)
```

NumPy arrays can be accessed using the same notation used for lists, but with an index for every dimension. 
<div class="alert alert-warning">
    
**Beware:** In NumPy arrays, the first index specifies the row (y-axis location), not the x-axis location, while the second index specifies the column (x-axis location), not the y-axis location.
</div>

```python
# Set the value at the Yth row, Xth column to 1
my_arr[3, 7] = 1
print(my_arr)
```

We can also use *slicing* `:` to determine whole regions of the array to read or assign. 

```python
# Insert a rectangle with the upper left corner at (Y_1,X_1) and the lower right corner at (Y_2,X_2)
my_arr[:3, 2:6] = 1
print(my_arr)
```

As you can see in the code, in slicing, we define an interval of numbers by `<lower_limit>:<upper_limit + 1>`. If one of the limits is not specified, it is interpreted as _from the beginning_ or _to the end_. Note that the lower limit is inclusive, while the upper limit is exclusive.

Using slicing we can also extract patches of an array to create new arrays. Run the example below and see its effect.

```python
# Let's extract the rectangle that we created
my_rect = my_arr[:3, 2:6]
# Note: '\n' creates a new line when printing
print(f'Original array: \n{my_arr}\n\nExtracted patch: \n{my_rect}')
```

Using NumPy arrays, we can easily perform the basic operations element-wise on all elements of one or several arrays (). Run the cell below to define a random array to work with.

```python
# First we define an array of size 5x5 filled with random integers from 0 to 10
np.random.seed(uid)
rand_array = np.random.randint(low=0, high=10, size=(5,5))
print(f'Initial array:\n{rand_array}')
```

We can perform any operation with a scalar for all elements of the array at the same time by using the array as if it was a variable with a single number. Run the example below and make sure you understand its results.

```python
# Add 2 to all elements
rand_array = rand_array + 2
print(f'rand_array + 2:\n{rand_array}\n')
# Divide all elements of the resulting array by 2
rand_array = rand_array / 2
print(f'rand_array / 2:\n{rand_array}')
```

Similarly, we can use expressions with multiple arrays of the same size and any of the basic operators: the calculation will be made on the co-positioned elements of the different arrays. Run the example below and make sure you understand its results.

```python
# Let's initialize two new random arrays
rand_array_1 = np.random.randint(low=0, high=10, size=(5,5))
rand_array_2 = np.random.randint(low=0, high=10, size=(5,5))
print(f'Array 1:\n{rand_array_1}\n\nArray 2:\n{rand_array_2}')
# We can for example add them together
added_arrays = rand_array_1 + rand_array_2
print(f'\nArray 1 + Array 2:\n{added_arrays}')
# Or multiply them
multiplied_arrays = rand_array_1 * rand_array_2
print(f'\nArray 1 * Array 2:\n{multiplied_arrays}')
```

We can also perform _advanced indexing_ of NumPy arrays: we can provide a [boolean](https://en.wikipedia.org/wiki/Boolean_data_type) array as the index and extract/modify only the values that are specified as `True`. An application of this form of indexing could be to take the multiplied array from the previous cell and set all its elements that are larger than $25$ to $100$. Run the example below and make sure you understand its results.

```python
# First print the multiplied array as a reference
print(f'Multiplied array:\n{multiplied_arrays}')
# Create the boolean array that contains true at every location, where the multiplied array has a value bigger than 25
boolean_array = multiplied_arrays > 25
print(f'\nBoolean array:\t(Multiplied array > 25)\n{boolean_array}')
# Now we can use this array to index the multiplied array and set all elements indicated by True to 100
multiplied_arrays[boolean_array] = 100
print(f'\nMultiplied array with all elements > 25 set to 100:\n{multiplied_arrays}')
```

Note that you can also create the boolean array in place, without explicitly declaring an additional variable. Run the cell below and see the difference with the preceeding cell. We will reuse the variable `multiplied_arrays`, and set all elements smaller than $20$ to $1$.

```python
# We get the array we have, and index only the elements smaller than 20.
multiplied_arrays[multiplied_arrays < 20] = 1
print(f'\nMultiplied array with all elements < 20 set to 1:\n{multiplied_arrays}')
```

(matplotlib)=
### 2.3. Matplotlib
[Back to table of contents](#Table-of-contents)

[Matplotlib](https://matplotlib.org/), specifically, [PyPlot module](https://matplotlib.org/stable/api/_as_gen/matplotlib.pyplot.html#module-matplotlib.pyplot) allows us to load images as NumPy arrays, and visualize NumPy arrays as images. 
 

It is best practice to import the module as `plt`.

```python
# Set the image display to be interactive 
%matplotlib widget
# Import pyplot
import matplotlib.pyplot as plt
```

The only functionality you really need to know is how to load images using `plt.imread`:

```python
# Load an example image located at 'images/epfl.tif'
example_img = plt.imread('images/epfl.tif')
# Then, we can check the image is really a NumPy array
type(example_img)
```

To see the image, use `plt.imshow`. In these labs, however, you will use the [imviewer](#4.-Python-Image-Viewer) class provided by us, which adds extra functionality on top of `plt.imshow`.

```python
# When using plt.imshow we first need to declare a new figure
plt.figure()
# Then we add the image to the figure
plt.imshow(example_img, cmap='gray')
# Give it a title
plt.title(f"Example Image - SCIPER: {int(getpass.getuser().split('-')[2]) if len(getpass.getuser().split('-')) > 2 else ord(getpass.getuser()[0])}")
# And finally we make sure it is displayed on the screen
plt.show()
```

Experiment with Python, NumPy and Matplotlib in the empty cell below. Learn to do the simple things that come to mind and take the chance to explore all of the options above. If you need some inspiration, run the next cell to follow the instructional video by Kay.

```python
from IPython.display import IFrame
IFrame(width="800", height="450", src="https://www.youtube-nocookie.com/embed/oewtZoQb2gY")
```

(python-image-viewer)=
# 3. Python Image Viewer
[Back to table of contents](#Table-of-contents)

To ease image visualization and manipulation, we have developed a library, the `imviewer` class, which you will use throughout the labs.
Before using the class, be sure to check the documentation [wiki](https://github.com/Biomedical-Imaging-Group/interactive-kit/wiki/Image-Viewer). 

The `imviewer` runs without explicitly importing any external libraries. However, we will import the library `Ipywidgets` ([see documentation](https://ipywidgets.readthedocs.io/en/latest/)) to allow for interactive visualization. You will even use it to add your own functions to the viewer's interface! 

How to import`imviewer`: 
```python
%matplotlib widget
from interactive_kit import imviewer as viewer
```
  
Run the cell below to import the viewer and load the images we will work with.

```python
# Configure plotting as dynamic
%matplotlib widget
# Import required packages for this exercise
import numpy as np
import matplotlib.pyplot as plt
import ipywidgets as widgets
# Import imviewer
from interactive_kit import imviewer as viewer

# Import images
car   = plt.imread('images/car_pad.tif')
hrct  = plt.imread('images/hrct.tif')
hands = plt.imread('images/hands.tif')
plate = plt.imread('images/plate.tif')
boats = plt.imread('images/boats.tif')
```

(creation-of-a-viewer)=
## 3.1. Creation of a viewer
[Back to table of contents](#Table-of-contents)

Run the cell below to follow an instructional video by Alejandro Noguerón Arámburu, one of the student developers. He will guide you and exemplify the use of the viewer.

```python
from IPython.display import IFrame
IFrame(width="800", height="450", src="https://www.youtube-nocookie.com/embed/OKYy3C0MCm0")
```

```python
# Declare image list
first_list = [hands, plate, car, hrct]
# Declare titles (only two)
title_list = ['Hands', 'Plate']
# First we close all matplotlib existing figures to improve performance
plt.close('all')
# Call the image viewer with the image list as first argument. The rest of the arguments are optional
first_viewer = viewer(first_list, title=title_list, widgets=True, hist=True, axis=True, cmap='gray')
```

To create an `imviewer` object, you simply call the function `viewer` and assign it a variable name. Ideally, the variable should be self-explanatory and end with the word *viewer* to avoid confusion (for example, `first_viewer` as above). Some of the most importante features of a `viewer` are:

**The colorbar option of the viewer is currently not working**

* **Images**: The same viewer can load many images at once. In this example, it was created with $4$ images. By default, you can switch between the images using the buttons `Prev` and `Next` on the widgets panel.
* **Titles**: You can give the viewer titles for each image, which will be shown immediately above it. If you do not provide a title for some images, the variable name will be shown by default. In this example, only $2$ were given.
* **Customization**: Several of the class' functionalities (e.g., `widgets`, `hist`, `axis`) can be enabled from the beginning. A user can later turn them off using the widgets or calling methods of the object.
* **Colormap**: The colormap to be used can be set from the start with the parameter `cmap = colormap`, where `colormap` is a string that specifies the desired colormap. It can be dynamically changed using `Options` and choosing one of the supported colormaps in the drop-down menu. 
* **Statistics**: On the lower part of the widget panel, the image statistics of the area of the image being displayed are shown. If you use the button with a little square at the left of the images to zoom into a selected region, you will see the statistics are updated in real time. You can also use the button with crossed double-arrows to pan through the image at a specific level of zoom. You can always use the button *Reset* to see the whole image again.

<div class="alert alert-warning">
    
<b>Beware: </b> The statistics displayed in the `imviewer` objects are rounded to the second decimal. If you need an exact statistic, use the functions `np.mean` or `np.std`on the images (NumPy arrays) directly.
</div>

* **Histogram**: Besides the normal *counts* vs *value* bars in a histogram, you will see a black line. This line represents *input* vs *output* intensity values. This line changes its position and slope if you change the brightness and contrast of the image. All pixel values that are to the left of the point where the line touches the bottom have the minimum intensity, and all that are to the right of the point where the line touches the top have the maximum intensity, and the ones that are in-between those two points have an intensity that is linearly scaled between the maximum and minumum value.

* **Display mode**: There are several options to display a list of images
    * Default display: Display only one image at a time, and use the buttons `Next` and `Prev` to browse the different images. 
    * Set `subplots=(m, n)`: Arrange the images in an $m \times n$ grid.

(using-widgets)=
## 3.2. Using widgets
[Back to table of contents](#Table-of-contents)

If you just want a quick visualization, you can call the `imviewer` with the image you want to display (or list of images). All the parameters will be set to their default values (`cmap = 'gray'`, single image mode, all the features set to off), but you can use the widgets to change this. Run the next cell to see all the default values. Go through all the widgets and explore their options while going through the explanatory cell below.

```python
# In order to keep the memory load on Noto low, we will generally close all figures before continuing.
# This means that the previous viewers are no longer interactive!
# But you can run their cell again if you need to.
plt.close('all')

# Declare image list
second_list = [hrct, hands, car]
# Call the image viewer with minimal arguments
second_viewer = viewer(second_list)
```

Your first view consists of:
* A **toolbar** with series of buttons at the left, which serve to control the dynamic environment. If you hover your mouse in the buttons, wou will see what each button does:
    * *Three horizontal lines*: Hide the toolbar
    * *House*: Reset the zoom of the images. 
    * *Cross with arrows*: Pan through an image. Click on it, then click on the image, hold and drag to pan through the axis.
    * *Square*: Zoom. Click on it, then click on the image, hold and drag to create a rectangular area. Release to zoom to the region.
    * *Floppy disk*: Take screenshot of figure and save in PNG format. 
* The **figure** holding the images (one figure is holding all the images). You can change the size of a figure using the small gray triangle (bottom-right corner). If you hover your cursor over the triangle, you will see it change to a two-sided arrow. By clicking on it you will be able to adjust the figure size.
* A **button** with the legend *Show Widgets*. Click on it and it will take you to the widget main menu. 

In the widgets main menu you will see the following widgets:
* **Brightness and Contrast** Slider: In this menu you will be able to change the color scaling of the image through a slider, given in percentages of the original maximum and minimum intensities.
* **Show Histogram** Button: Show or hide the histograms of the images.
* **Options** Button: It will take you to the options menu, where you will be able to show or hide the axis and change the colormap.
* **Reset** Button: This will reset the all the parameters to the default state (colormap, brightness/contrast ect.), while the _House Button_ on the top left only resets the display area of the viewer and doesn't change any parameters.
* **Stats** Textbox: Here you will see the mean, the standard deviation, the range of values and the size of the image. Mean, standard deviation and range of values will be updated when you zoom to a region. 

### Multiple choice question

Answer the following question using the viewer that you just defined.

In the image `car`, what is the mean value of **the carpet that the car is standing on**? 
Hint: Zoom into several regions that are not too close to the edge nor to the car.

1. Around 40
2. Around 60 
3. Around 80

In the following cell, modify the variable `answer` to your actual answer, e.g., `answer=1`, `answer=2`, or `answer=3`.

<!-- grade_id: cell-40970ade3d3b9d3e_2022, points: 0, grade: false, solution: true -->
```python
# Assign your answer to this variable
answer = None

# YOUR CODE HERE
```

<!-- grade_id: cell-9db1b31c6c9f4ebb_2022, points: 1, grade: true, solution: false -->
```python
# Let's do a sanity check - check that the answer is one of the possible choices
assert answer in [1, 2, 3], "Select one of the answers: 1, 2 or 3."

```

(user-defined-widgets)=
## 3.3. User defined widgets
[Back to table of contents](#Table-of-contents)

`imviewer` allows you to create a function in a Jupyter notebook, and apply it simultaneously to all images within your `imviewer` object for a set of slider values. The function you create will take as parameters:
* an image (`NumPy array`), and 
* one or more numerical values.


In the next cells, you will see a very basic example. We will apply a pixel-wise operation on the image `car`: all the pixels with a value below a treshold (given in $\%$) will be set to the *maximum* value. 

First, we define the thresholding function `your_function`. Run the next cell to define it.

```python
# Define your function
def your_function(image, threshold):    
    # We make a copy of the original, where we will apply our threshold.
    output = np.copy(image)
    # For greater flexibility, we will get the value as a percentage of the maximum
    value = threshold * 0.01 * np.amax(image)
    # Apply threshold
    output[image < value] = np.amax(image)
    return output
```

Besides defining your function, you need to define the corresponding slider(s), a button to run `your_function` with the sliders' values, and the specific function that will run `your_function` i.e., the slider's _callback_ function.

<div class = "alert alert-info">
    
<b>Note</b>: Although this is rather abstract, the example will make it clear. We will use an integer slider in the range $[0,100]$ (`widgets.IntSlider`), a button that clearly specifies what our function does (`widgets.Button`), and a callback function `callback` that calls `your_function` with the slider's value as threshold. Run the next cell, click the button *Extra Widgets*, and try different values. Remember to click on `Apply Threshold`!
</div>
    
If you want some guidance, run the next cell to follow an instructional video by Alejandro.

```python
from IPython.display import IFrame
IFrame(width="800", height="450", src="https://www.youtube-nocookie.com/embed/__JXopJPLxQ")
```

```python
# close all open figures
plt.close('all')

# Declare slider
threshold_slider = widgets.IntSlider(value=0, min=0, max=100, step=1, 
                                     description='Threshold', continuous_update=False)

# Declare button with meaningful description
activation_button = widgets.Button(description='Apply Threshold')

# Declare callback
def callback(img):
    # Get slider value
    threshold = threshold_slider.value
    # Call your function
    output = your_function(img, threshold)
    return output

# Call viewer, passing the widget and callback separately as lists
threshold_viewer = viewer(car, new_widgets=[threshold_slider, activation_button], callbacks=[callback], widgets=True)
```

(programmatic-customization)=
## 3.4. Programmatic customization
[Back to table of contents](#Table-of-contents)

Most options can be modified either through widgets or in code. Run the following cells. In the first one, we will create the object `car_viewer`, which displays the image `car`. By running the next cells, one by one, we will produce changes on the `car_viewer` without interacting with it directly. 

```python
plt.close('all')
car_viewer = viewer(car)
```

```python
car_viewer.set_axis(axis=True)
```

```python
car_viewer.set_widgets(widgets=True)
```

```python
car_viewer.set_colormap(colormap='viridis')
```

```python
car_viewer.show_histogram(hist=True)
```

```python
car_viewer.set_axis(axis=False)
```

(3.5-try-it-yourself)=
## 3.5. Try it yourself!
[Back to table of contents](#Table-of-contents)

### Creating an imviewer
Create in the next cell a viewer object with the following characteristics (do it with code, on one or several lines):
* The first image should be `hrct` and the second `plate`,
* The titles are left as default,
* The colormap is set to `'inferno'`,
* The axis are set to `False`,
* Only a single image at a time is displayed.

Store your object in the variable `exercise_viewer`.

<!-- grade_id: cell-7c287f72627b42ee_2022, points: 0, grade: false, solution: true -->
```python
plt.close('all')
exercise_viewer = None
# YOUR CODE HERE
```

If you want to check the state of your `imviewer` object, you can check [the documentation](https://github.com/Biomedical-Imaging-Group/interactive-kit/wiki/Image-Viewer) and explore the attributes. See the two examples in the next cell. Both should show that there are indeed two images. The second one is in the form of an assertion (it will produce an error if the condition is not met), to let you know if your viewer is in the right direction.

<!-- grade_id: cell-18c524bc087e7f47_2022, points: 1, grade: true, solution: false -->
```python
# Sanity check
assert exercise_viewer.number_images == 2, f'Your viewer should have 2 images! Instead it has {exercise_viewer.number_images}.'
print('Well done! Your viewer seems to be working.')
```

### Threshold operation 

Now you will program a function and add a slider to apply this function to the image (see Section [4.C.](#4.C.-User-defined-widgets)).

Your first task is to define a function `range_function(img, rmin, rmax)` that will only **keep** pixel values inside a range of values $[r_{\mathrm{min}}, r_{\mathrm{max}}]$ ($r_{\mathrm{min}}$ and $r_{\mathrm{max}}$ included), while the values outside this range will be set to $0$.
You will apply it on the image `car`, so `rmin` and `rmax` should be between $0$ and $255$. Define this function in the following cell. 

<div class="alert alert-info">
  
<b>Note: </b> Use boolean indexing of the NumPy array `img`, as in the thresholding example, to perform the necessary operations.
</div>

<!-- grade_id: cell-f7d3dd947484c706_2022, points: 0, grade: false, solution: true -->
```python
# Define your function
def range_function(img, rmin, rmax):    
    # Best practices are creating a copy. You don't wanna mess the original.
    output = np.copy(img)
    # set the pixels of output that are smaller than rmin or bigger than rmax to 0
    # YOUR CODE HERE
    return output
```

You can test your function with the following cell. In it, we define a linear array `test_array` that is formed by the integers from 0 to 255 in ascending order, using the method [NumPy.arange](https://NumPy.org/doc/stable/reference/generated/NumPy.arange.html). Then, we apply your function with the range $[100, 199]$ to it. There should be exactly $100$ elements that were not set to $0$. If the cell runs smoothly, your implementation is probably correct.
<div class="alert alert-info">
    
**Note:** The function `np.count_nonzero(arr)` counts the number of elements in `arr` that are not zero.
</div>

<!-- grade_id: cell-d32bcff0667c88a9_2022, points: 1, grade: true, solution: false -->
```python
test_array = np.arange(0, 256)
assert np.count_nonzero(range_function(test_array, 100, 199)) == 100, 'Your implementation is not yet correct.'
print('Congratulations! Your function is correct.')

```

Declare the necessary widgets to control this function (slider(s) and a button, as well as the button's callback function). Use one `IntRangeSlider` ([see doc](https://ipywidgets.readthedocs.io/en/latest/examples/Widget%20List.html#IntRangeSlider)). When you get the value of this slider (by calling `IntRangeSlider.value`), you get a list with two elements corresponding to the minimum and the maximum value of the range. You can use these values to call your function (`range_slider.value[0]` and `range_slider.value[1]` if `range_slider` has been defined as an `IntRangeSlider` widget). 

Use the next cell to declare your widgets (store the slider in the variable `range_slider`, your button on `range_button` and your callback on `range_callback(img)`)

<!-- grade_id: cell-a4cde69b86b888bf_2022, points: 0, grade: false, solution: true -->
```python
# Declare the slider 
range_slider = None
# YOUR CODE HERE

# Declare the button
range_button = None
# YOUR CODE HERE

# Declare the callback
def range_callback(img):
    # Define output
    output = None
    
    # Get slider value
    # YOUR CODE HERE
    
    # Call your function, e.g., output = range_function(img, ...)
    # YOUR CODE HERE
    
    return output
```

Now we will create an `imviewer` using your widgets. Explore the results of your function (click the button *Extra Widgets*), explore different values and see if you implemented it correctly.

<!-- grade_id: cell-fa45b07997e9a3be_2022, points: 1, grade: true, solution: false -->
```python
# Call viewer, passing the widget and callback separately as lists, e.g., range_viewer = viewer(car, ...)
plt.close('all')
range_viewer = viewer(car, new_widgets=[range_slider, range_button], callbacks=[range_callback], widgets=True)
```

(image-processing-in-python)=
# 4. Image processing in Python 
[Back to table of contents](#Table-of-contents)

Throughout this course we will focus on one most common image processing library
[**`scikit-image`**](https://scikit-image.org/).

Scikit stands for SciPy Toolkits, it is developed specifically for image-processing tasks. These toolkits are add-on packages for SciPy. They have become very well known and have a large community of users and developers. 

For a short instructional video by Alejandro and Stan on how to check the documentation this library run the cell below.

```python
from IPython.display import IFrame
IFrame(width="800", height="450", src="https://www.youtube-nocookie.com/embed/e1rsEUgDDko")
```

Just as with `NumPy` (`np`) or `matplotlib.pyplot` (`plt`), there is a standard abbreviation to import this library. 

```python
# Scikit-image
import skimage
```

(examples)=
# 4.1. Examples
[Back to table of contents](#Table-of-contents)

Now we will go through two very basic operations: rotation and edge detection for you to get familiar with scikit-image. We will show you the basic syntax. For the full list of parameters and options, go to the official documentation of each function. 


### Rotation in Scikit Image

The basic syntax of the function for grayscale images is  (see [the documentation](https://scikit-image.org/docs/stable/api/skimage.transform.html?highlight=rotate#skimage.transform.rotate)):
```python
output = skimage.transform.rotate(image, angle, resize=False, center=None, mode='reflect', preserve_range=True)
```

The parameters are:
* `input`: The source image, 
* `angle`: Rotation angle in degrees,
* `resize`: Specify whether to resize the image to contain the full original. Defaults to `False`,
* `center`: Specifies the center of rotation. The order is (columns, rows),
* `mode`: Specify boundary conditions (see the last example). 


We will apply sciki-image's rotation to `car` for an angle of $60^\circ$ with the center of rotation at $(50, 50)$. Run the next cell to see the result. Modify the cell and experiment with the different parameters.

```python
# First, we import the module transform
from skimage import transform

# Now we apply operation
rot_car_skimage = skimage.transform.rotate(car, angle=60, resize=False, center=(50, 50), mode='reflect', preserve_range=True)

# And visualize
plt.close('all')
skimage_rotate_viewer = viewer([car, rot_car_skimage], title=['Original', 'Rotated'], subplots=(1,2))
```

### Edge Detection in Scikit Image

The basic syntax of the Sobel filter for grayscale images is  (see [the documentation](https://scikit-image.org/docs/dev/api/skimage.filters.html#skimage.filters.sobel)):

```python
output = skimage.filters.sobel(image, axis=None, mode='reflect')
```

The parameters are:
* `image`: The source image, 
* `axis`: If given, compute the axis in the specified direction. If `None`, the magnitude of the gradient is calculated,
* `mode`: Specify boundary conditions (see example above). 


Run the next cell to apply Scikit-Image Sobel filter on the image `plate`. Modify the cell and play with the different values and with different images.

```python
# First, we import the module filters
from skimage import filters

# Apply filter
plate_edges_skimage = skimage.filters.sobel(plate, axis=None, mode='reflect')

# Visualize
plt.close('all')
skimage_edge_viewer = viewer([plate, plate_edges_skimage], title=['Original', 'Edges'], subplots=(1,2))
```

(4.2-try-it-yourself)= 
## 4.2. Try it yourself! 
[Back to table of contents](#Table-of-contents)

Complete the function `gaussian_blur`, which implements a **Gaussian smoothing** on the given input image. Set
* the standard deviation of the Gaussian kernel to be $5$,
* boundary conditions to be reflective, and
* keep the original range of values.

<div class='alert alert-warning'>
    <b>Note:</b> In python it's important to always specify the parameter names when using a function. Only this way, you can be sure that you set the correct parameters. For example, don't use <code>skimage.filters.sobel(plate, None, 'reflect')</code>, but instead use <code>skimage.filters.sobel(plate, axis=None, mode='reflect')</code>.
</div>

<!-- grade_id: cell-3aad12f740231a32_2022, points: 0, grade: false, solution: true -->
```python
# Function that applies a Gaussian blur
def gaussian_blur(img):
    out = np.zeros(img.shape)
    
    # YOUR CODE HERE
    
    return out
```

<!-- grade_id: cell-1919408b3ba0b1c7_2022, points: 2, grade: true, solution: false -->
```python
# Visualize 
plt.close('all')
gaussian_viewer = viewer([boats, gaussian_blur(boats)], title=['Original', 'Blurred'], subplots=(1,2))
```

<div class="alert alert-success">

Congratulations on finishing this first lab of IP1!
</div>