```{image} https://www.epfl.ch/about/overview/wp-content/uploads/2020/07/logo-epfl-1024x576.png 
:width: 140px
:align: left
```
## Image Processing Laboratory Notebooks
---

This Juypter notebook is part of a series of computer laboratories which are designed
to teach image-processing programming; they are running on the EPFL's Noto server. They are the practical complement of the theoretical lectures of the EPFL's Master course 
[**MICRO-511 Image Processing I**](https://moodle.epfl.ch/course/view.php?id=522) taught by Prof. M. Unser and Prof. D. Van de Ville.

The project is funded by the Center for Digital Education and the School of Engineering. It is owned by the [Biomedical Imaging Group](http://bigwww.epfl.ch/). 
The distribution or the reproduction of the notebook is strictly prohibited without the written consent of the authors.  &copy; EPFL 2023.

**Authors**: 
    [Pol del Aguila Pla](mailto:pol.delaguilapla@epfl.ch), 
    [Kay Lächler](mailto:kay.lachler@epfl.ch),
    [Alejandro Noguerón Arámburu](mailto:alejandro.nogueronaramburu@epfl.ch),
    [Yan Liu](mailto:yan.liu@epfl.ch), and
    [Daniel Sage](mailto:daniel.sage@epfl.ch).
    
---
# Lab 0: Introduction
**Released**: Thursday September 28, 2023

Double-click on this cell, fill your name and SCIPER number below to verify your identity in Noto and set the seed for random results.
:::{attention} Please write down your name and SCIPER! 
### Student Name: 
### SCIPER: 
:::

<!-- grade_id: cell-136865f50e6e2d19_2022, points: 0, grade: true, solution: false -->
```sos
%use sos
import getpass
# This line recovers your camipro number to mark the images with your ID
uid = int(getpass.getuser().split('-')[2]) if len(getpass.getuser().split('-')) > 2 else ord(getpass.getuser()[0])
print(f'SCIPER: {uid}')
```

# Introduction

This introductory lab will guide you through the tools that you will need for the labs. 


(toc)=
## Table of contents

1. [Introduction to Jupyter notebooks](#intro-to-notebooks) 
    1. [Cellular structure](#cellular-structure)
    2. [The SoS kernel](#sos)
2. [Python and JavaScript](#python-js)
    1. [Basic Python programming](#basic-python)
        1. [NumPy](#np)
        2. [Matplotlib](#mpl)
    2. [Basic Javascript programming](#basic-js)
    3. [Python vs JavaScript speed comparison](#python-js-speed)
3. [Javascript image-access class](#image-access) 
4. [Python Image Viewer](#python-image-viewer)
    1. [Creation of a viewer](#creation-of-a-viewer)
    2. [Using widgets](#using-widgets) 
    3. [User-defined widgets](#user-defined-widgets)
    4. [Programmatic customization](#programmatic-customization)
    5. [Try it yourself!](#4.5-try-it-yourself) 
5. [Image processing in Python](#image-processing-in-python)
    1. [Examples](#examples)
    2. [Try it yourself!](#5.2-try-it-yourself) 

(intro-to-notebooks)=
# 1. Introduction to Jupyter notebooks
(cellular-structure)=
## 1.A. Cellular structure
[Back to table of contents](#toc)

A Jupyter notebook consists of cells that contain either code snippets with their results or text. The cellular structure of Jupyter notebooks allows us to run small snippets of code one after the other, go back and forth between, or check intermediate results between them. 

:::{note} If you declare some variable in a cell, this variable will be available in all other cells that use the same programming language, irrespective of the location of the cell. 
:::
Let's look at an example. Read the description of the cells below and run them as instructed.

In this first cell we declare a variable `a` and assign it the value $5$. Run it.

```sos
%use sos
# cell 1
# this is our variable
a = 5
```

In the second cell we print the value of `a`. If you run it just after the cell above, it should print "The value of a is 5". Run it.

```sos
%use sos
# cell 2
# here we print the value of a
print(f"The value of a is {a}")
```

Then, in the third cell, we modify `a` by adding $2$ to it. Run it.

```sos
%use sos
# cell 3
# add two to a
a = a + 2
```

Now run _cell 2_ again: the value of `a` has been modified, even though _cell 2_ is located above _cell 3_ in our notebook. Now run _cell 1_ again, where we set `a=5`, and `a` will be $5$ in all cells again. 
This shows the **global** impact every single cell has on the notebook. 

:::{note} If at any point the kernel should crash for whatever reason, you can restart it by clicking on the button `Kernel` in the toolbar on top and then click on `Restart Kernel...`. Once the kernel is restarted, you will need to run all necessary cells again to regain the working environment you had before. A good practice to test that there are no errors in the notebook is to click `Restart and Run All Cells...` once you have finished a notebook. This will run all the cells from top to bottom and only stop if an error occurs.
:::

(sos)=
## 1.B. The SoS kernel
[Back to table of contents](#toc)

The [SoS](https://vatlab.github.io/sos-docs/) kernel allows for programming with different languages. It is indicated by the line `%use sos` at the beginning of each cell which you should not delete or modify.

You only need to know that:
- `%use sos` indicates that a cell is written in Python,
- `%use javascript` indicates that a cell is written in JavaScript,
- `%put var_name ` converts the variable `var_name` from JavaScript to Python
- `%get var_name ` converts the variable `var_name` from Python to JavaScript

The keywords above are called _magics_ and they always need to be on the first lines of code with nothing in-between them. 

Follow an instructional video below by Kay Lächler, one of the student developers.

:::{iframe} https://www.youtube-nocookie.com/embed/5XPVnavkgyY
:width: 80%
:::

First, we declare some variable `my_var` in python.

```sos
%use sos
# This is our variable, let's initialize it to 0
my_var = 0
```

Next, we want to modify this variable in JavaScript, so we define a JavaScript cell to do that. Run the next cell **without modifying it**, and see if you understand the resulting error.

```sos
%use javascript
// %get my_var 
// %put my_var

// try to modify the python variable without converting it to JavaScript first
my_var = my_var + 5
```

Now, let's do it properly. Uncomment (remove the  `//`) the second and the third lines. These lines import the variable  from python using `%get my_var `.
Once all the code in the cell has been executed, `my_var` will be convert it back to python using `%put my_var `. 

If we now print the value of `my_var` in Python, it should be 5.

```sos
%use sos
# Print the value of the python variable my_var
print(f'The value of my_var is {my_var}')
```

(python-js)=
# 2. Python and JavaScript
[Back to table of contents](#toc)

_Python_ is a high-level language with many great libraries with built-in image processing algorithms. You will learn how to apply the different image processing methods to images using the fastest and most standard tools and visualize the results. 
_JavaScript_ is a fairly low-level programming language, fit for implementing image processing algorithms by accessing each pixel in an image. You will understand how image processing methods work in detail, pixel by pixel.


In these notebooks, you will usually implement every method **once** in JavaScript to understand how it works. Then, apply the same method in Python using existing  libraries. In many cases, you will also design workflows for different image processing applications in Python.

    
:::{note}If you are already familiar with [Python](python.org) and the [NumPy](numpy.org) and [Matplotlib](matplotlib.org) libraries, you can jump to Section [2.B.](#2.B.-Basic-JavaScript-programming).
:::

(basic-python)=
## 2.A. Basic Python programming
[Back to table of contents](#toc)

There are three basic types of variables  in Python:
```python
# Numbers
my_var = 2
my_var = 5.2
# Strings
my_var = 'Either use single quotes'
my_var = "or use double quotes"
# Boolean
my_var = True
my_var = False
```
Every variable can hold every type of data, so we don't need to specify the type at initialization.

We can have lists of several items:
```python
# List of numbers
my_list = [2.5, 7, 3.3]
# List of strings
my_list = ['E.T.', 'phone', 'home']
# List of booleans
my_list = [True, True, False]
```
Lists are 0-indexed and can be accessed with square brackets:
```python
# Set the 3rd element of my_list to 5
my_list[2] = 5
# Retrieve the 2nd element of my_list
new_var = my_list[1]
```
Similarly we can make tuples, which are just like lists, but their values cannot be modified after they have been initialized:
```python
# This is a tuple of three numbers
my_tuple = (3, 8.6, 4)
# Access to a tuple is similar to lists
new_var = my_tuple[1]
# But we cannot write to a tuple!
```
For numbers we have the following basic operators:
```python
# Define the two variables
a = 5
b = 2
# Addition
a + b
>> 7
# Subtraction
a - b
>> 3
# Multiplication
a * b
>> 10
# Division
a / b
>> 2.5
# Integer division
a // b
>> 2
# Modulo (remainder of integer division)
a % b
>> 1
# Power
a ** b
>> 25
```
We can have conditional statements
```python
if a < b:
    # printing is done with the print function
    print('a is smaller than b')
else:
    print('a is bigger than or equal to b')
```
<div class="alert alert-info">

:::{attention} Indentation is important in Python. Writing a conditional statement, a loop, a function, etc., without an indent (`Tab`) raises an error.
:::

with the following comparison operators:
```python
<  # Smaller than
<= # Smaller or equal
>  # Bigger than
>= # Bigger than or equal to
== # Equal to
!= # Not equal to
```
A `for` loop is declared as:
```python
# Note: range(n) creates numbers from 0 to n-1
for i in range(10):
    # Another for loop inside the for loop - we call this 'nested for loops'
    for j in range(10):
        # We can also have formatted strings indicated by a preceeding 'f', where the number i is inserted into the string
        print(f'Loop iteration number {j}')
```

Similarly a `while` loop:
```python
i = 0
while i < 10:
    print(f'Loop iteration number {i}')
    i = i + 1
```
And finally, a function can be defined using the `def` keyword:
```python
# A function that adds two numbers and returns the result
def addition(a, b):
    # Add the numbers
    result = a + b
    # Return the result
    return result

# Run the function
addition(5, 10)
>> 15
```

Experiment with Python in the cell below. Learn to do the simple things that come to mind and take the chance to explore all of the options above.

```sos
%use sos

```

(np)=
### 2.A.a. NumPy
[Back to table of contents](#toc)

[NumPy](https://NumPy.org/) library provides multidimensional matrix/array processing tools. Most image processing libraries in Python use NumPy arrays to represent images. 

It is best practice to import the NumPy library as `np`, to avoid typing `numpy` every time. This also makes the code more readable. Run the cell below to import NumPy.

```sos
%use sos
# Import the NumPy library
import numpy as np
```

To create a new NumPy array filled with zeros we can use `np.zeros`. It requires the size of the array we want to get as a tuple `(height, width)`. Run the cell below to see an example.

Optionally, we can specify the internal numerical [type](https://NumPy.org/devdocs/user/basics.types.html) of the array. For example, one can use 8-bit unsigned integers (`np.uint8`). This type is convenient for 8-bit images, as the resulting array will occupy less memory than with other types, such as `np.float32` or `np.float64`. For most of the labs, you will not need to mind the numerical type of your variables: it adapts automatically to represent what you are storing in each instance.

```sos
%use sos
# A NumPy array of size 5x10 with numerical type "unsigned 8-bit integer"
my_arr = np.zeros((5,10), dtype=np.uint8)
# Print the result
print(my_arr)
```

NumPy arrays can be accessed using the same notation used for lists, but with an index for every dimension. 
<div class="alert alert-warning">
    
:::{attention} In NumPy arrays, the first index specifies the row (y-axis location), not the x-axis location, while the second index specifies the column (x-axis location), not the y-axis location.
:::

```sos
%use sos
# Set the value at the Yth row, Xth column to 1
my_arr[3, 7] = 1
print(my_arr)
```

We can also use *slicing* `:` to determine whole regions of the array to read or assign. 

```sos
%use sos
# Insert a rectangle with the upper left corner at (Y_1,X_1) and the lower right corner at (Y_2,X_2)
my_arr[:3, 2:6] = 1
print(my_arr)
```

As you can see in the code, in slicing, we define an interval of numbers by `<lower_limit>:<upper_limit + 1>`. If one of the limits is not specified, it is interpreted as _from the beginning_ or _to the end_. Note that the lower limit is inclusive, while the upper limit is exclusive.

Using slicing we can also extract patches of an array to create new arrays. Run the example below and see its effect.

```sos
%use sos
# Let's extract the rectangle that we created
my_rect = my_arr[:3, 2:6]
# Note: '\n' creates a new line when printing
print(f'Original array: \n{my_arr}\n\nExtracted patch: \n{my_rect}')
```

Using NumPy arrays, we can easily perform the basic operations element-wise on all elements of one or several arrays (). Run the cell below to define a random array to work with.

```sos
%use sos
# First we define an array of size 5x5 filled with random integers from 0 to 10
np.random.seed(uid)
rand_array = np.random.randint(low=0, high=10, size=(5,5))
print(f'Initial array:\n{rand_array}')
```

We can perform any operation with a scalar for all elements of the array at the same time by using the array as if it was a variable with a single number. Run the example below and make sure you understand its results.

```sos
%use sos
# Add 2 to all elements
rand_array = rand_array + 2
print(f'rand_array + 2:\n{rand_array}\n')
# Divide all elements of the resulting array by 2
rand_array = rand_array / 2
print(f'rand_array / 2:\n{rand_array}')
```

Similarly, we can use expressions with multiple arrays of the same size and any of the basic operators: the calculation will be made on the co-positioned elements of the different arrays. Run the example below and make sure you understand its results.

```sos
%use sos
# Let's initialize two new random arrays
rand_array_1 = np.random.randint(low=0, high=10, size=(5,5))
rand_array_2 = np.random.randint(low=0, high=10, size=(5,5))
print(f'Array 1:\n{rand_array_1}\n\nArray 2:\n{rand_array_2}')
# We can for example add them together
added_arrays = rand_array_1 + rand_array_2
print(f'\nArray 1 + Array 2:\n{added_arrays}')
# Or multiply them
multiplied_arrays = rand_array_1 * rand_array_2
print(f'\nArray 1 * Array 2:\n{multiplied_arrays}')
```

We can also perform _advanced indexing_ of NumPy arrays: we can provide a [boolean](https://en.wikipedia.org/wiki/Boolean_data_type) array as the index and extract/modify only the values that are specified as `True`. An application of this form of indexing could be to take the multiplied array from the previous cell and set all its elements that are larger than $25$ to $100$. Run the example below and make sure you understand its results.

```sos
%use sos
# First print the multiplied array as a reference
print(f'Multiplied array:\n{multiplied_arrays}')
# Create the boolean array that contains true at every location, where the multiplied array has a value bigger than 25
boolean_array = multiplied_arrays > 25
print(f'\nBoolean array:\t(Multiplied array > 25)\n{boolean_array}')
# Now we can use this array to index the multiplied array and set all elements indicated by True to 100
multiplied_arrays[boolean_array] = 100
print(f'\nMultiplied array with all elements > 25 set to 100:\n{multiplied_arrays}')
```

Note that you can also create the boolean array in place, without explicitly declaring an additional variable. Run the cell below and see the difference with the preceeding cell. We will reuse the variable `multiplied_arrays`, and set all elements smaller than $20$ to $1$.

```sos
%use sos
# We get the array we have, and index only the elements smaller than 20.
multiplied_arrays[multiplied_arrays < 20] = 1
print(f'\nMultiplied array with all elements < 20 set to 1:\n{multiplied_arrays}')
```

(mpl)=
### 2.A.b. Matplotlib
[Back to table of contents](#toc)

[Matplotlib](https://matplotlib.org/), specifically, [PyPlot module](https://matplotlib.org/stable/api/_as_gen/matplotlib.pyplot.html#module-matplotlib.pyplot) allows us to load images as NumPy arrays, and visualize NumPy arrays as images. 
 

It is best practice to import the module as `plt`.

```sos
%use sos
# Set the image display to be interactive 
%matplotlib widget
# Import pyplot
import matplotlib.pyplot as plt
```

The only functionality you really need to know is how to load images using `plt.imread`:

```sos
%use sos
# Load an example image located at 'images/epfl.tif'
example_img = plt.imread('images/epfl.tif')
# Then, we can check the image is really a NumPy array
type(example_img)
```

To see the image, use `plt.imshow`. In these labs, however, you will use the [imviewer](#4.-Python-Image-Viewer) class provided by us, which adds extra functionality on top of `plt.imshow`.

```sos
%use sos
# When using plt.imshow we first need to declare a new figure
plt.figure()
# Then we add the image to the figure
plt.imshow(example_img, cmap='gray')
# Give it a title
plt.title(f"Example Image - SCIPER: {int(getpass.getuser().split('-')[2]) if len(getpass.getuser().split('-')) > 2 else ord(getpass.getuser()[0])}")
# And finally we make sure it is displayed on the screen
plt.show()
```

Experiment with Python, NumPy and Matplotlib in the empty cell below. Learn to do the simple things that come to mind and take the chance to explore all of the options above. If you need some inspiration, follow the instructional video below by Kay.

:::{iframe} https://www.youtube-nocookie.com/embed/oewtZoQb2gY
:width: 80%
:::

```sos
%use sos

```

(basic-js)=
## 2.B. Basic JavaScript programming
[Back to table of contents](#toc)

If you are already familiar with JavaScript, you can jump to Section [2.C.](#2.C.-Python-and-JavaScript-speed-comparison)


To declare a variable in JavaScript one uses the `var` statement. There are three basic types of variables:
```javascript
// numbers
var my_var = 2;
var my_var = 5.2;
// strings
var my_var = 'Either use single quotes';
var my_var = "or use double quotes";
// boolean
var my_var = true;
var my_var = false;
```
Every variable can hold every type of data, so we don't need to specify the type at initialization.
<div class="alert alert-info">

:::{note} In JavaScript, every line should be terminated with a semicolon `;`.
:::

We can have arrays of several items. There are two equivalent ways of constructing an array, the first one is:
```javascript
// list of numbers
var my_array = [2.5, 7, 3.3];
// list of strings
var my_array = ['E.T.', 'phone', 'home'];
// list of booleans
var my_array = [true, true, false];
```

And the second one is:
```javascript
// list of numbers
var my_array = new Array(2.5, 7, 3.3);
// list of strings
var my_array = new Array('E.T.', 'phone', 'home');
// list of booleans
var my_array = new Array(true, true, false);
```



:::{note} With the second approach, you are explicitly calling the _constructor_ of the class Array (in other words, the function that creates new Arrays). Consequently, you need to use the keyword `new`. We will see this also in Section [3](#3.-Javascript-image-access-class) with the class we have prepared for you to process images simply with JavaScript.
:::

Arrays are 0-indexed and can be accessed with square brackets:
```javascript
// set the 3rd element of my_list to 5
my_array[2] = 5;
// retrieve the 2nd element of my_list
var new_var = my_array[1];
```
Contrary to Python, tuples do not exist in JavaScript.

For numbers we have the following basic operators:
```javascript
// define the two variables
var a = 5;
var b = 2;
// addition
a + b;
>> 7
// subtraction
a - b;
>> 3
// multiplication
a * b;
>> 10
// division
a / b;
>> 2.5
// integer division can be emulated through
parseInt(a / b);
>> 2
// modulo (remainder of the integer division)
a % b;
>> 1
// power
a ** b;
>> 25
```
We can have conditional statements
```javascript
if(a < b){
    // printing in JS is done with console.log
    console.log('a is smaller than b');
}else{
    console.log('a is bigger than b');
}
```

    
:::{attention}Indentation is not important in JavaScript but curly brackets `{}` are. Yet as good practice, we still recommend you to indent properly as it makes your code easier to read. To write a conditional statement, a loop, a function, etc., use curly brackets to define the beginning and end of each block of code. 
:::

with the following comparison operators:
```javascript
<  // smaller than
<= // smaller or equal
>  // bigger than
>= // bigger or equal
== // equal
!= // not equal
```
A `for` loop is declared as:
```javascript
for(var i = 0; i < 10; i++){
    // another for loop inside the for loop
    for(var j = 0; j < 10; i++){
        // we can concatenate strings and numbers using the + operator
        console.log('Loop iteration number ' + j);
    }
}
```
Similarly a `while` loop:
```javascript
var i = 0;
while(i < 10){
    console.log('Loop iteration number ' + i);
    i++; // increment i by one
}
```

And finally, a function can be defined using the `function` keyword:
```javascript
// a function that adds two numbers and returns the result
function addition(a, b){
    // Add the numbers
    var result = a + b;
    // Return the result
    return result;
}

// run the function
addition(5, 10);
>> 15
```

Experiment with JavaScript in the cells below. Learn to do the simple things that come to mind and take the chance to explore all of the options above. If you want to visualize what you try out, use the cell further below, Python, and SoS to play around.

```sos
%use javascript
```

```sos
%use sos
```

(python-js-speed)=
## 2.C. Python vs JavaScript speed comparison
[Back to table of contents](#toc)

As mentioned before, JavaScript will mainly be used to implement pixel-level algorithms, which require nested `for` loops. To illustrate why this is necessary, we'll compare the time it takes for both Python and JavaScript to perform a zero-padded convolution implemented using nested `for` loops.

The two-dimensional convolution between two images `h` and `f` is given by
$$(h \ast f)[k,l] = \sum_{m \in \mathbb{Z}}\sum_{n \in \mathbb{Z}}f[m,n]h[k-m,l-n]$$
Don't worry, you don't need to know how that works right now, you'll see this in Chapter _3.2_ of the course. But it's safe to say that the convolution is a very important operation for image processing, which is why we'll use it for this demonstration.

First let's define a Python function to generate a Gaussian kernel, which we will use as one of the images for the convolution. Run the cell below to define this function and obtain and visualize the Gaussian impulse response.

```sos
%use sos
# Imports and definitions, in case you skipped the previous parts
%matplotlib widget
import numpy as np
import matplotlib.pyplot as plt
import getpass

# Function that creates a Gaussian convolutional kernel based on the input parameters
def gaussian_kernel(sig, length):
    # Generate the 1D grid-points
    ax = np.linspace(-(length - 1) / 2., (length - 1) / 2., length)
    # Generate the 2D grid-points created by ax x ax
    xx, yy = np.meshgrid(ax, ax)
    # Calculate the Gaussian kernel (up to proportionality constant)
    kernel = np.exp(-0.5 * (np.square(xx) + np.square(yy)) / np.square(sig))
    # Normalize the kernel to sum 1 and return it
    return kernel / np.sum(kernel)

# Get a Gaussian kernel
ker = gaussian_kernel(sig=3, length=5)
# Show the result
plt.figure()
plt.imshow(ker, cmap='gray')
plt.title(f"Gaussian kernel - SCIPER: {int(getpass.getuser().split('-')[2]) if len(getpass.getuser().split('-')) > 2 else ord(getpass.getuser()[0])}")
plt.show()
```

From the implementation above you can learn how to use NumPy functions such as [`np.linspace`](https://numpy.org/doc/stable/reference/generated/numpy.linspace.html), [`np.meshgrid`](https://numpy.org/doc/stable/reference/generated/numpy.meshgrid.html), [`np.exp`](https://numpy.org/doc/stable/reference/generated/numpy.exp.html), [`np.square`](https://numpy.org/doc/stable/reference/generated/numpy.square.html) and [`np.sum`](https://numpy.org/doc/stable/reference/generated/numpy.sum.html) (click on any of the links to go to the documentation). Try them out at the end of the Python section if you are curious! Otherwise, don't worry, this example will always be here for you.

Now that we have the Gaussian kernel ready, let's define the JavaScript convolution function. Run the cell below to define it.

```sos
%use javascript
// function that convolves img with kernel
function convolution(img, kernel){
    // initialize the output image as a copy of the input image
    var smoothed_img = [...img];
    // compute the kernel offset
    const k_offset = Math.floor(kernel.length / 2);
    const l_offset = Math.floor(kernel[0].length / 2);
    // perform zero-padded convolution
    for(i = 0; i < img.length; i++){
        for(j = 0; j < img[0].length; j++){
            // compute output for pixel i, j. 
            // zero-padding: extend the image with zeros beyond its boundaries.
            smoothed_img[i][j] = 0;
            for(k = 0; k < kernel.length; k++){
                for(l = 0; l < kernel[0].length; l++){
                    // compute evaluation points to enforce boundary
                    var x = k - k_offset + i;
                    var y = l - l_offset + j;
                    // increment if appropriate (if outside the image, don't increment, there are zeros)
                    if((x < img.length) && (x > 0) && (y < img[0].length) && (y > 0)){
                        smoothed_img[i][j] += img[x][y] * kernel[k][l]; 
                    } 
                }
            }
        }
    }
    // return the output image
    return smoothed_img;
}
```

From the implementation above, you can learn useful JavaScript functions like `Math.floor`. Nonetheless, some of the tricks we had to use will not be necessary for you. For example, the lines
```javascript
var smoothed_img = [...img];
k < kernel.length;
l < kernel[0].length;
```
will be much more intuitive to implement using our class [image-access](#3.-Javascript-image-access-class) below, so do not worry about them. 

Now, let's define the Python convolution function. Run the cell below to define it.

```sos
%use sos
# Function that convolves img with kernel
def convolution(img, kernel):
    # Initialize output image as an empty image
    smoothed_img = np.zeros(img.shape)
    # Compute the kernel offset
    k_offset = kernel.shape[0] // 2
    l_offset = kernel.shape[1] // 2
    # Perform zero-padded convolution
    for i in range(img.shape[0]):
        for j in range(img.shape[1]):
            # Compute output for pixel i, j. 
            # Zero-padding: extend the image with zeros beyond its boundaries.
            for k in range(kernel.shape[0]):
                for l in range(kernel.shape[1]):
                    # Compute evaluation points to enforce boundary
                    x = k - k_offset + i
                    y = l - l_offset + j
                    # Increment if appropriate (if outside the image, don't increment, there are zeros)
                    if (x < img.shape[0]) and (x > 0) and (y < img.shape[1]) and (y > 0):
                        smoothed_img[i, j] += img[x, y] * kernel[k, l]
    # Return the output image
    return smoothed_img
```

From the implementation above, you can learn useful NumPy functions like [`np.zeros`](https://numpy.org/doc/stable/reference/generated/numpy.zeros.html), or properties like [`img.shape`](https://numpy.org/doc/stable/reference/generated/numpy.ndarray.shape.html), when `img` is a NumPy array. As always, feel free to try them out now at the end of the Python section!

Now, we load an image and transfer it to JavaScript, together with the Gaussian kernel we defined before in Python. In this way, we will be able to test the JavaScript function we implemented. Run the cell below to load the image, display it, and transfer the image and the Gaussian kernel to Javascript.

```sos
%use sos
%put hrct --to javascript
%put ker --to javascript
# Load the hrct image
hrct = plt.imread('images/hrct.tif')
# Display it
plt.figure()
plt.imshow(hrct, cmap='gray')
plt.title(f"HRCT - SCIPER: {int(getpass.getuser().split('-')[2]) if len(getpass.getuser().split('-')) > 2 else ord(getpass.getuser()[0])}")
plt.show()
```

Now, we are ready to run the convolution in JavaScript and measure the time it takes. Run the cell below to do it.

```sos
%use javascript
%put convolved_hrct_js
%put js_time 
// get the starting time in miliseconds
var start = new Date();

// run the convolution
var convolved_hrct_js = convolution(hrct, ker);

// get the ending time in miliseconds
var end = new Date();

// calculate the elapsed time in miliseconds
var js_time = end - start

// display the elapsed time
console.log('The JavaScript convolution took ' + js_time + ' miliseconds to run.');
```

Now, we do the same for the implementation of the convolution in Python. Run the cell below to do it.

:::{hint}You can go grab a drink while this is running... As you may have noticed, in a Jupyter Notebook a cell is still running if inside the square brackets in the top left corner of the cell you see an asterisk `[*]` instead of a number.
:::

:::{note}
To measure the time in Python we will import the `time` module.
:::

```sos
%use sos
import time

# Get starting time in seconds
start = time.time()

# Run the convolution
convolved_hrct_jpython = convolution(hrct, ker)

# Get the ending time seconds
end = time.time()

# Calculate the elapsed time in miliseconds
python_time = round((end - start)*1000)

# Display the elapsed time
print(f'The Python convolution took {python_time} miliseconds to run.');
print(f'That makes JavaScript around {round(python_time/js_time)} times faster than Python for this implementation.')
```

Hopefully, you can understand from this illustration that standard Python is not made for pixel-level algorithms with nested `for` loops. This justifies the use of JavaScript for such implementations. Just so you can actually see the result of the function we just ran, run the next cell to see it.

:::{attention} Languages that allow [**vectorization**](https://en.wikipedia.org/wiki/Automatic_vectorization) like Python (through the NumPy library) and MATLAB allow operations to be performed on whole arrays. This is much simpler to code and as fast as a `for` loop in a low-level programming language. That is exactly the point of using Python, and in general, **using `for` loops to iterate through images will not be allowed during the labs**, even if your result is correct.
:::

```sos
%use sos
plt.figure()
plt.imshow(convolved_hrct_jpython, cmap='gray')
plt.title(f"HRCT convolved with a Gaussian kernel - SCIPER: {int(getpass.getuser().split('-')[2]) if len(getpass.getuser().split('-')) > 2 else ord(getpass.getuser()[0])}")
plt.show()
```

(image-access)=
# 3. Javascript image-access class
[Back to table of contents](#toc)

While JavaScript offers the big advantage of fast nested `for` loops for pixel-by-pixel programming, it also comes with a downside. It does not offer convenient array/image processing libraries like Python. To make things easier, we created our own JavaScript class, [`image-access`](https://github.com/Biomedical-Imaging-Group/image-access/wiki).
Every `image-access` object is a container for an image and provides easy-to-use methods for pixel access that take care of boundary conditions. 

In all upcoming labs, we will import the `image-access` class as `Image`. Run the JavaScript cell below to do it.

<!-- grade_id: cell-6e236bcfb6aba184_2022, points: 0, grade: true, solution: false -->
```sos
%use javascript
var Image = require('image-access')
```

An `Image` object of dimensions `height` $\times$ `width` can be created with
```javascript
var img = new Image(height, width);
```

Optionally, `height` and `width` can be provided as a single array `[height, width]`. This makes it easy to create an image with the same dimensions as an existing image by using the `shape` method, for example
```javascript
var img = new Image(an_existing_image.shape());
```

By default, the image is initialized to $0$ for every pixel, unless specified by the `init_value` attribute of the `options` object

```javascript
var options = {};
options.init_value = 255;
var img = new Image(height, width, options);
```

If we want to create an RGB image, we specify the `rgb` attribute in the `options` object to `true`
```javascript
var options = {};
options.rgb = true;
var img = new Image(height, width, options);
```

An `Image` object can also be created from an existing JavaScript array, by simply passing it as a single argument
```javascript
var img = new Image(an_existing_array);
```
This is useful, for example, to create an `Image` object from a Python NumPy array converted to Javascript.

We now look at how to use the basic functionalities in the class. The most important operations are the `img.setPixel` and `img.getPixel` methods, exactly as their names suggest

* The `img.setPixel(x, y, value)` method can be used to set the pixel at location `(x,y)` to a specific `value`:
```javascript
// here we set the pixel (3,5) of img to 255
img.setPixel(3, 5, 255);
```
* The `img.getPixel(x, y)` method can be used to retrieve the value of the pixel at location `(x,y)`:
```javascript
// here we retrieve the pixel (3,5) of img
var pixel = img.getPixel(3, 5);
```
By default, the `setPixel` and `getPixel` methods will use _mirroring_ boundary conditions to provide values for pixels outside the image range. 

The height and width of an `Image` object can be retrieved either using using the `img.shape()` method, which returns a list `[height, width]` or, more conveniently, with `img.nx` and `img.ny`.
```javascript
// retrieve height and width of the image as a list
height_and_width_list = img.shape();
// get only the width
width = img.nx;
// get only the height
height = img.ny;
```
This second option is very useful when setting the limits in `for` loops. <br><br>

We can extract the neighbourhood of size `width` $\times$ `height` around the pixel at location `(x,y)`, using the `img.getNbh(x, y, width, height)` method:
```javascript
// extracting the 3x3 neighbourhood around pixel (12,7)
nbh = img.getNbh(12, 7, 3, 3);
```
By default, the method will use _mirroring_ boundary conditions, just like the `getPixel` and `setPixel` methods.

Finally, if we need to convert an `Image` object to a simple JavaScript array, we can do it with the `img.toArray()` method, for example
```javascript
// convert the Image object to a JavaScript array
img_array = img.toArray();
```
This will be necessary, for example, to convert the image to Python.

Create a new `Image` object called `img`, of size $50 \times 50$, and set the pixels at locations $(x=10, y=20)$, $(x=15, y=30)$ and $(x=40, y=7)$ to $255$.

<!-- grade_id: cell-779a0057e28178dd_2022, points: 0, grade: false, solution: true -->
```sos
%use javascript
%put img 

// create a new Image object named img of size 50 x 50
// YOUR CODE HERE

// set the pixel at location (x=10,y=20) to 255
// YOUR CODE HERE
// set the pixel at location (x=15,y=30) to 255
// YOUR CODE HERE
// set the pixel at location (x=40,y=7) to 255
// YOUR CODE HERE
```

Run the cell below to (partially) check if you did the right thing.

<!-- grade_id: cell-e5d66a3787cd6867_2022, points: 2, grade: true, solution: false -->
```sos
%use javascript
if((img.nx != 50) || (img.ny != 50)){
    throw new Error("The image size is not correct.");
}else{console.log('Good job! Your image has the correct size.')}

if(img.getPixel(40, 7) != 255){
    throw new Error("The pixel at location (x=40, y=7) should have value 255.")
}else{console.log('Good job! The pixel at (x=40, y=7) was properly set.')}

```

Experiment with the `image-access` class in the empty cell below. Learn to do the simple things that come to mind and take the chance to explore all of the options above. If you need some inspiration, follow the instructional video below  by one of our TAs Stan.

:::{iframe} https://www.youtube-nocookie.com/embed/03BrgZlBIsE
:width: 80%
:::

```sos
%use javascript
var I = new Image(5,5);
console.log(I.getNbh(2,2,3,3).visualize())
```

(python-image-viewer)=
# 4. Python Image Viewer
[Back to table of contents](#toc)

To ease image visualization and manipulation, we have developed a library, the `imviewer` class, which you will use throughout the labs.
Before using the class, be sure to check the documentation [wiki](https://github.com/Biomedical-Imaging-Group/interactive-kit/wiki/Image-Viewer). 

The `imviewer` runs without explicitly importing any external libraries. However, we will import the library `Ipywidgets` ([see documentation](https://ipywidgets.readthedocs.io/en/latest/)) to allow for interactive visualization. You will even use it to add your own functions to the viewer's interface! 

How to import`imviewer`: 
```python
%matplotlib widget
from interactive_kit import imviewer as viewer
```
  
Run the cell below to import the viewer and load the images we will work with.

```sos
%use sos
# Configure plotting as dynamic
%matplotlib widget
# Import required packages for this exercise
import numpy as np
import matplotlib.pyplot as plt
import ipywidgets as widgets
# Import imviewer
from interactive_kit import imviewer as viewer

# Import images
car   = plt.imread('images/car_pad.tif')
hrct  = plt.imread('images/hrct.tif')
hands = plt.imread('images/hands.tif')
plate = plt.imread('images/plate.tif')
boats = plt.imread('images/boats.tif')
```

(creation-of-a-viewer)=
## 4.A. Creation of a viewer
[Back to table of contents](#toc)

Follow an instructional video below by Alejandro Noguerón Arámburu, one of the student developers. He will guide you and exemplify the use of the viewer.

:::{iframe} https://www.youtube-nocookie.com/embed/OKYy3C0MCm0
:width: 80%
:::

```sos
%use sos

# Declare image list
first_list = [hands, plate, car, hrct]
# Declare titles (only two)
title_list = ['Hands', 'Plate']
# First we close all matplotlib existing figures to improve performance
plt.close('all')
# Call the image viewer with the image list as first argument. The rest of the arguments are optional
first_viewer = viewer(first_list, title=title_list, widgets=True, hist=True, axis=True, cmap='spring')
```

To create an `imviewer` object, you simply call the function `viewer` and assign it a variable name. Ideally, the variable should be self-explanatory and end with the word *viewer* to avoid confusion (for example, `first_viewer` as above). Some of the most importante features of a `viewer` are:

**The colorbar option of the viewer is currently not working**

* **Images**: The same viewer can load many images at once. In this example, it was created with $4$ images. By default, you can switch between the images using the buttons `Prev` and `Next` on the widgets panel.
* **Titles**: You can give the viewer titles for each image, which will be shown immediately above it. If you do not provide a title for some images, the variable name will be shown by default. In this example, only $2$ were given.
* **Customization**: Several of the class' functionalities (e.g., `widgets`, `hist`, `axis`) can be enabled from the beginning. A user can later turn them off using the widgets or [calling methods](#prog_customization) of the object. 
* **Colormap**: The colormap to be used can be set from the start with the parameter `cmap = colormap`, where `colormap` is a string that specifies the desired colormap. It can be dynamically changed using `Options` and choosing one of the supported colormaps in the drop-down menu. 
* **Statistics**: On the lower part of the widget panel, the image statistics of the area of the image being displayed are shown. If you use the button with a little square at the left of the images to zoom into a selected region, you will see the statistics are updated in real time. You can also use the button with crossed double-arrows to pan through the image at a specific level of zoom. You can always use the button *Reset* to see the whole image again.

:::{warning} The statistics displayed in the `imviewer` objects are rounded to the second decimal. If you need an exact statistic, use the functions [`np.mean`](https://numpy.org/doc/stable/reference/generated/numpy.mean.html) or [`np.std`](https://numpy.org/doc/stable/reference/generated/numpy.std.html) on the images (NumPy arrays) directly.
:::

* **Histogram**: Besides the normal *counts* vs *value* bars in a histogram, you will see a black line. This line represents *input* vs *output* intensity values. This line changes its position and slope if you change the brightness and contrast of the image. All pixel values that are to the left of the point where the line touches the bottom have the minimum intensity, and all that are to the right of the point where the line touches the top have the maximum intensity, and the ones that are in-between those two points have an intensity that is linearly scaled between the maximum and minumum value.

* **Display mode**: There are several options to display a list of images
    * Default display: Display only one image at a time, and use the buttons `Next` and `Prev` to browse the different images. 
    * Set `subplots=(m, n)`: Arrange the images in an $m \times n$ grid.

(using-widgets)=
## 4.B. Using widgets
[Back to table of contents](#toc)

If you just want a quick visualization, you can call the `imviewer` with the image you want to display (or list of images). All the parameters will be set to their default values (`cmap = 'gray'`, single image mode, all the features set to off), but you can use the widgets to change this. Run the next cell to see all the default values. Go through all the widgets and explore their options while going through the explanatory cell below.

```sos
%use sos
# In order to keep the memory load on Noto low, we will generally close all figures before continuing.
# This means that the previous viewers are no longer interactive!
# But you can run their cell again if you need to.
plt.close('all')

# Declare image list
second_list = [hrct, hands, car]
# Call the image viewer with minimal arguments
second_viewer = viewer(second_list)
```

Your first view consists of:
* A **toolbar** with series of buttons at the left, which serve to control the dynamic environment. If you hover your mouse in the buttons, wou will see what each button does:
    * *Three horizontal lines*: Hide the toolbar
    * *House*: Reset the zoom of the images. 
    * *Cross with arrows*: Pan through an image. Click on it, then click on the image, hold and drag to pan through the axis.
    * *Square*: Zoom. Click on it, then click on the image, hold and drag to create a rectangular area. Release to zoom to the region.
    * *Floppy disk*: Take screenshot of figure and save in PNG format. 
* The **figure** holding the images (one figure is holding all the images). You can change the size of a figure using the small gray triangle (bottom-right corner). If you hover your cursor over the triangle, you will see it change to a two-sided arrow. By clicking on it you will be able to adjust the figure size.
* A **button** with the legend *Show Widgets*. Click on it and it will take you to the widget main menu. 

In the widgets main menu you will see the following widgets:
* **Brightness and Contrast** Slider: In this menu you will be able to change the color scaling of the image through a slider, given in percentages of the original maximum and minimum intensities.
* **Show Histogram** Button: Show or hide the histograms of the images.
* **Options** Button: It will take you to the options menu, where you will be able to show or hide the axis and change the colormap.
* **Reset** Button: This will reset the all the parameters to the default state (colormap, brightness/contrast ect.), while the _House Button_ on the top left only resets the display area of the viewer and doesn't change any parameters.
* **Stats** Textbox: Here you will see the mean, the standard deviation, the range of values and the size of the image. Mean, standard deviation and range of values will be updated when you zoom to a region. 

### Multiple choice question

Answer the following question using the viewer that you just defined.

In the image `car`, what is the mean value of **the carpet that the car is standing on**? 
Hint: Zoom into several regions that are not too close to the edge nor to the car.

1. Around 40
2. Around 60 
3. Around 80

In the following cell, modify the variable `answer` to your actual answer, e.g., `answer=1`, `answer=2`, or `answer=3`.

<!-- grade_id: cell-40970ade3d3b9d3e_2022, points: 0, grade: false, solution: true -->
```sos
%use sos
# Assign your answer to this variable
answer = None

# YOUR CODE HERE
```

<!-- grade_id: cell-9db1b31c6c9f4ebb_2022, points: 1, grade: true, solution: false -->
```sos
%use sos
# Let's do a sanity check - check that the answer is one of the possible choices
assert answer in [1, 2, 3], "Select one of the answers: 1, 2 or 3."

```

(user-defined-widgets)=
## 4.C. User defined widgets
[Back to table of contents](#toc)

`imviewer` allows you to create a function in a Jupyter notebook, and apply it simultaneously to all images within your `imviewer` object for a set of slider values. The function you create will take as parameters:
* an image (`NumPy array`), and 
* one or more numerical values.


In the next cells, you will see a very basic example. We will apply a pixel-wise operation on the image `car`: all the pixels with a value below a treshold (given in $\%$) will be set to the *maximum* value. 

First, we define the thresholding function `your_function`. Run the next cell to define it.

```sos
%use sos
# Define your function
def your_function(image, threshold):    
    # We make a copy of the original, where we will apply our threshold.
    output = np.copy(image)
    # For greater flexibility, we will get the value as a percentage of the maximum
    value = threshold * 0.01 * np.amax(image)
    # Apply threshold
    output[image < value] = np.amax(image)
    return output
```

Besides defining your function, you need to define the corresponding slider(s), a button to run `your_function` with the sliders' values, and the specific function that will run `your_function` i.e., the slider's _callback_ function.

:::{note} Although this is rather abstract, the example will make it clear. We will use an integer slider in the range $[0,100]$ `widgets.IntSlider`, a button that clearly specifies what our function does `widgets.Button`, and a callback function `callback` that calls `your_function` with the slider's value as threshold. Run the next cell, click the button *Extra Widgets*, and try different values. Remember to click on `Apply Threshold`!
:::
    
If you want some guidance, follow an instructional video below by Alejandro.

:::{iframe} https://www.youtube-nocookie.com/embed/__JXopJPLxQ
:width: 80%
:::

```sos
%use sos
# close all open figures
plt.close('all')

# Declare slider
threshold_slider = widgets.IntSlider(value=0, min=0, max=100, step=1, 
                                     description='Threshold', continuous_update=False)

# Declare button with meaningful description
activation_button = widgets.Button(description='Apply Threshold')

# Declare callback
def callback(img):
    # Get slider value
    threshold = threshold_slider.value
    # Call your function
    output = your_function(img, threshold)
    return output

# Call viewer, passing the widget and callback separately as lists
threshold_viewer = viewer(car, new_widgets=[threshold_slider, activation_button], callbacks=[callback], widgets=True)
```

(programmatic-customization)=
## 4.D. Programmatic customization
[Back to table of contents](#toc)

Most options can be modified either through widgets or in code. Run the following cells. In the first one, we will create the object `car_viewer`, which displays the image `car`. By running the next cells, one by one, we will produce changes on the `car_viewer` without interacting with it directly. 

```sos
%use sos
plt.close('all')
car_viewer = viewer(car)
```

```sos
%use sos
car_viewer.set_axis(axis=True)
```

```sos
%use sos
car_viewer.set_widgets(widgets=True)
```

```sos
%use sos
car_viewer.set_colormap(colormap='viridis')
```

```sos
%use sos
car_viewer.show_histogram(hist=True)
```

```sos
%use sos
car_viewer.set_axis(axis=False)
```

(4.5-try-it-yourself)=
## 4.E. Try it yourself!
[Back to table of contents](#toc)

### Creating an imviewer
Create in the next cell a viewer object with the following characteristics (do it with code, on one or several lines):
* The first image should be `hrct` and the second `plate`,
* The titles are left as default,
* The colormap is set to `'inferno'`,
* The axis are set to `False`,
* Only a single image at a time is displayed.

Store your object in the variable `exercise_viewer`.

<!-- grade_id: cell-7c287f72627b42ee_2022, points: 0, grade: false, solution: true -->
```sos
%use sos
plt.close('all')
exercise_viewer = None
# YOUR CODE HERE
```

If you want to check the state of your `imviewer` object, you can check [the documentation](https://github.com/Biomedical-Imaging-Group/interactive-kit/wiki/Image-Viewer) and explore the attributes. See the two examples in the next cell. Both should show that there are indeed two images. The second one is in the form of an assertion (it will produce an error if the condition is not met), to let you know if your viewer is in the right direction.

<!-- grade_id: cell-18c524bc087e7f47_2022, points: 1, grade: true, solution: false -->
```sos
%use sos
# Sanity check
assert exercise_viewer.number_images == 2, f'Your viewer should have 2 images! Instead it has {exercise_viewer.number_images}.'
print('Well done! Your viewer seems to be working.')
```

### Threshold operation

Now you will program a function and add a slider to apply this function to the image (see Section [4.C.](#4.C.-User-defined-widgets)).

Your first task is to define a function `range_function(img, rmin, rmax)` that will only **keep** pixel values inside a range of values $[r_{\mathrm{min}}, r_{\mathrm{max}}]$ ($r_{\mathrm{min}}$ and $r_{\mathrm{max}}$ included), while the values outside this range will be set to $0$.
You will apply it on the image `car`, so `rmin` and `rmax` should be between $0$ and $255$. Define this function in the following cell. 

<div class="alert alert-info">
  
<b>Note: </b> Use boolean indexing of the NumPy array `img`, as in the thresholding example, to perform the necessary operations.
</div>

<!-- grade_id: cell-f7d3dd947484c706_2022, points: 0, grade: false, solution: true -->
```sos
%use sos
# Define your function
def range_function(img, rmin, rmax):    
    # Best practices are creating a copy. You don't wanna mess the original.
    output = np.copy(img)
    # set the pixels of output that are smaller than rmin or bigger than rmax to 0
    # YOUR CODE HERE
    return output
```

You can test your function with the following cell. In it, we define a linear array `test_array` that is formed by the integers from 0 to 255 in ascending order, using the method [NumPy.arange](https://NumPy.org/doc/stable/reference/generated/NumPy.arange.html). Then, we apply your function with the range $[100, 199]$ to it. There should be exactly $100$ elements that were not set to $0$. If the cell runs smoothly, your implementation is probably correct.
:::{note} The function `np.count_nonzero(arr)` counts the number of elements in `arr` that are not zero.
:::

<!-- grade_id: cell-d32bcff0667c88a9_2022, points: 1, grade: true, solution: false -->
```sos
%use sos
test_array = np.arange(0, 256)
assert np.count_nonzero(range_function(test_array, 100, 199)) == 100, 'Your implementation is not yet correct.'
print('Congratulations! Your function is correct.')

```

Declare the necessary widgets to control this function (slider(s) and a button, as well as the button's callback function). Use one `IntRangeSlider` ([see doc](https://ipywidgets.readthedocs.io/en/latest/examples/Widget%20List.html#IntRangeSlider)). When you get the value of this slider (by calling `IntRangeSlider.value`), you get a list with two elements corresponding to the minimum and the maximum value of the range. You can use these values to call your function (`range_slider.value[0]` and `range_slider.value[1]` if `range_slider` has been defined as an `IntRangeSlider` widget). 

Use the next cell to declare your widgets (store the slider in the variable `range_slider`, your button on `range_button` and your callback on `range_callback(img)`)

<!-- grade_id: cell-a4cde69b86b888bf_2022, points: 0, grade: false, solution: true -->
```sos
%use sos
# Declare the slider 
range_slider = None
# YOUR CODE HERE

# Declare the button
range_button = None
# YOUR CODE HERE

# Declare the callback
def range_callback(img):
    # Define output
    output = None
    
    # Get slider value
    # YOUR CODE HERE
    
    # Call your function, e.g., output = range_function(img, ...)
    # YOUR CODE HERE
    
    return output
```

Now we will create an `imviewer` using your widgets. Explore the results of your function (click the button *Extra Widgets*), explore different values and see if you implemented it correctly.

<!-- grade_id: cell-fa45b07997e9a3be_2022, points: 1, grade: true, solution: false -->
```sos
%use sos
# Call viewer, passing the widget and callback separately as lists, e.g., range_viewer = viewer(car, ...)
plt.close('all')
range_viewer = viewer(car, new_widgets=[range_slider, range_button], callbacks=[range_callback], widgets=True)
```

(image-processing-in-python)=
# 5. Image processing in Python
[Back to table of contents](#toc)

Throughout this course we will focus on one most common image processing library
[**`scikit-image`**](https://scikit-image.org/).

Scikit stands for SciPy Toolkits, it is developed specifically for image-processing tasks. These toolkits are add-on packages for SciPy. They have become very well known and have a large community of users and developers. 

Follow a short instructional video below by Alejandro and Stan on how to check the documentation this library.

:::{iframe} https://www.youtube-nocookie.com/embed/e1rsEUgDDko
:width: 80%
:::

Just as with `NumPy` (`np`) or `matplotlib.pyplot` (`plt`), there is a standard abbreviation to import this library. 

```sos
%use sos
# Scikit-image
import skimage
```

(examples)=
# 5.A. Examples
[Back to table of contents](#toc)

Now we will go through two very basic operations: rotation and edge detection for you to get familiar with scikit-image. We will show you the basic syntax. For the full list of parameters and options, go to the official documentation of each function. 


### Rotation in Scikit Image

The basic syntax of the function for grayscale images is  (see [the documentation](https://scikit-image.org/docs/stable/api/skimage.transform.html?highlight=rotate#skimage.transform.rotate)):
```python
output = skimage.transform.rotate(image, angle, resize=False, center=None, mode='reflect', preserve_range=True)
```

The parameters are:
* `input`: The source image, 
* `angle`: Rotation angle in degrees,
* `resize`: Specify whether to resize the image to contain the full original. Defaults to `False`,
* `center`: Specifies the center of rotation. The order is (columns, rows),
* `mode`: Specify boundary conditions (see the last example). 


We will apply sciki-image's rotation to `car` for an angle of $60^\circ$ with the center of rotation at $(50, 50)$. Run the next cell to see the result. Modify the cell and experiment with the different parameters.

```sos
%use sos
# First, we import the module transform
from skimage import transform

# Now we apply operation
rot_car_skimage = skimage.transform.rotate(car, angle=60, resize=False, center=(50, 50), mode='reflect', preserve_range=True)

# And visualize
plt.close('all')
skimage_rotate_viewer = viewer([car, rot_car_skimage], title=['Original', 'Rotated'], subplots=(1,2))
```

### Edge Detection in Scikit Image

The basic syntax of the Sobel filter for grayscale images is  (see [the documentation](https://scikit-image.org/docs/dev/api/skimage.filters.html#skimage.filters.sobel)):

```python
output = skimage.filters.sobel(image, axis=None, mode='reflect')
```

The parameters are:
* `image`: The source image, 
* `axis`: If given, compute the axis in the specified direction. If `None`, the magnitude of the gradient is calculated,
* `mode`: Specify boundary conditions (see example above). 


Run the next cell to apply Scikit-Image Sobel filter on the image `plate`. Modify the cell and play with the different values and with different images.

```sos
%use sos
# First, we import the module filters
from skimage import filters

# Apply filter
plate_edges_skimage = skimage.filters.sobel(plate, axis=None, mode='reflect')

# Visualize
plt.close('all')
skimage_edge_viewer = viewer([plate, plate_edges_skimage], title=['Original', 'Edges'], subplots=(1,2))
```

(5.2-try-it-yourself)=
## 5.B. Try it yourself!
[Back to table of contents](#toc)

Complete the function `gaussian_blur`, which implements a **Gaussian smoothing** on the given input image. Set
* the standard deviation of the Gaussian kernel to be $5$,
* boundary conditions to be reflective, and
* keep the original range of values.

:::{attention} In python it's important to always specify the parameter names when using a function. Only this way, you can be sure that you set the correct parameters. For example, don't use `skimage.filters.sobel(plate, None, 'reflect')`, but instead use `skimage.filters.sobel(plate, axis=None, mode='reflect')`.
:::

<!-- grade_id: cell-3aad12f740231a32_2022, points: 0, grade: false, solution: true -->
```sos
%use sos
# Function that applies a Gaussian blur
def gaussian_blur(img):
    out = np.zeros(img.shape)
    
    # YOUR CODE HERE
    
    return out
```

<!-- grade_id: cell-1919408b3ba0b1c7_2022, points: 2, grade: true, solution: false -->
```sos
%use sos
# Visualize 
plt.close('all')
gaussian_viewer = viewer([boats, gaussian_blur(boats)], title=['Original', 'Blurred'], subplots=(1,2))
```

Congratulations on finishing this first lab of IP1!


:::{attention}
Make sure to save your notebook (you might want to keep a copy on your personal computer) and submit it on Moodle before the deadline. Check the [Moodle](https://moodle.epfl.ch/course/view.php?id=522) page of the course for detailed instructions of submission. **Please do not rename the notebook!**
:::